# Phys 260: Python assignment header

### (1) Fill out the cell below.  
The cell below is a **code cell**.  Fill out your University of Michigan uniqname, then your name, and collaborators in the cell below **inside the quotes**.  

**Do not delete the quotes.**  We will use this information to organize your assignments.  To edit and execute cells, double click inside the cell, type, and press \<shift\>+\<enter\> to execute.

```python
UNIQNAME = ""
NAME = ""
COLLABORATORS = ""
```

### (2) Check your python version.  
**Execute the cell below** (double click in the cell and press \<shift\>+\<enter\>, or click in the cell and press the Run button) to check that you are using a version of python that is compatible with the tool we are using to grade your assignments.  If your ```IPython``` version is too old, we will *not* be able to grade your assignments.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

### (3) Do your best to answer all questions in the assignment.  
To answer questions, **replace** anything that says either
- "YOUR ANSWER HERE" 
- 
```
YOUR CODE HERE
raise NotImplementedError
``` 

with your answer/code.  Cells with either of the two bullet points above are cells of the notebook that will be graded.

**To edit markdown** cells (e.g. this one),  *double click in the cell to type*.  Press \<shift\>+\<enter\> to execute the cell.  Try editing the text below to replace the with your information (e.g. Camille Avestruz, cavestru):  

[first name] [last name], uniqname


### (4) Make sure your notebook runs sequentially.
After you complete this assignment, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

```python

```

# Introduction -- Reminder

Each Python lab will start with a pre-flight exercise that walks through building some of the set up and tools ($\sim$ 30 min), followed by an in-class tutorial with time for Q+A (50 min) so you can walk through steps that will be necessary for the homework assignment you will submit ($\sim$ 3 hrs).  Each lab will contain starter code, similar to what you see below.  Please fill in the code to complete the pre-flight assignment in preparation for the in-class tutorial.  

Preflight ($\sim$30-60 min, 10 points) **Typically due: Wednesdays 3pm EST**

*Preflight typically graded by Wednesday 5p EST*

In-class tutorial and Q+A ($\sim$ 50 min, 10 points) **Typically occurs: Thursdays 9am EST**

Homework assignment ($\sim$ 3-5 hrs, 30 points) **Typically due: Mondays 9am EST** 

*Homework typically graded by Thursday 5p*

When we grade your homework, we will not run your code. Once submitted, your notebook should have the outputs for all of your results.  Please do not include long outputs from debugging, beyond a few print statements and the requested visualizations (i.e. plots).

**Grading:** When we grade your notebook, we will convert the .ipynb file to an HTML file.  We will be using [nbgrader](https://nbgrader.readthedocs.io/en/stable/) to grade your notebooks.  **Note:** If your ```IPython``` version is too old, we will *not* be able to grade your assignments.


# Phys 260 Python Tutorial/HW 6: Simulating RC circuits (30 points total)

## Tutorial/HW summary
- Setting up the RC circuit with different initial conditions (with/without a battery)
- Solving a more complex circuit

```python
import numpy as np
from matplotlib import pyplot as plt
```

## RC circuit to model

We will model an RC circuit with the analogous steps as the `Spring` from the preflight.  Here, we will examine the behavior of:

- Case 1: An uncharged capacitor in series with a resistor and a battery that will be charged from the battery.  
- Case 2: A discharging capacitor as current flows through a resistor after the battery has been disconnected and the circuit closed.

In the first case, the inital voltage across the capacitor is then zero, and the initial current through the resistor will be $I = \frac {V_b} {R}$ where $V_b$ is the voltage of the battery.  The schematic is below:

<img src="http://www-personal.umich.edu/~gerdes/img/MC-PythonCircuit02.png" height="400" width="400">


We can write down the differential equation which includes this initial condition, and solve numerically.  Plot the current through the resistor and the charge on the capacitor as a function of time.  Following this, we can compute and plot the power delivered by the battery to the circuit as a function of time.  We can also compute the total energy delivered by the battery to the circuit, compare this to the energy stored in the capacitor energy dissipated as heat in the resistor? 

For the first set up, we will keep the loop rule in mind,
\begin{eqnarray}
V - IR - Q/C = 0
\end{eqnarray}


### Practice writing equations in latex (3 points - 7 min)
In the cell below, write: 
- the current at the ith timestep, $I_i=\frac{\Delta Q_i}{\Delta t}$, in terms of everything else (you'll need to rearrange the loop rule), and 
- the charge on the capacitor at the next timestep, $Q_{i+1}$, in terms of $Q_{i}$, where $Q_{i}$ is the charge on the capacitor at the ith timestep (which determines the current at the ith timestep). 
- the current in the i+1 timestep, $I_{i+1}$, given the new charge, $Q_{i+1}$, on the capacitor.

<!-- grade_id: cell-59cd3722d249db02, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Build a class to model the RC circuit (2 points)

Below, we have a class to model the RC circuit.  
- Fill out the internal method, `self._set_i_now` to set the attribute `self.i_now` to the appropriate value dependent on other attributes.
- Use the internal method where needed.

<!-- grade_id: cell-142a0a81f661f167, points: 0, grade: false, solution: true -->
```python
class RCCircuit :
    """Creates an instance of an rc circuit model, which can be time evolved"""
    def __init__(self, resistance, capacitance, q_capacitor, voltage_battery) :
        """A model of an rc circuit
        Parameters
        ----------
        resistance : float
            value of resistance in ohms
        capacitance : float 
            value of capacitance in Farads
        q_capacitor : float
            initial charge on capacitor
        voltage_battery : float
            voltage across battery
        """

        self.resistance = resistance
        self.capacitance = capacitance
        
        # The time constant is a characteristic timescale of the system, analogous to the period in a spring
        self.rc_constant = resistance * capacitance
        
        self.voltage_battery = voltage_battery
        self.q_now = q_capacitor
        
        #  Use the internal method _set_i_now
        # YOUR CODE HERE
        raise NotImplementedError()
        
        self.i_vals = []
        self.q_vals = []

    def _set_i_now(self) :
        """Calculate current for attribute i_now in terms of other attributes"""
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def _set_num_timesteps(self, num_timescales, num_timesteps_per_timescale) :        
        """Internal method. Sets the attributes num_timesteps and dt.
        Parameters
        ----------
        num_timescales : floattimestep 
            number of timescales to evolve over
        num_timesteps_per_timescale : int
            number of timesteps per characteristic timescale, defines the time resolution 
        """
        
        dt = self.rc_constant / num_timesteps_per_timescale

        try :
            if self.dt != dt :
                # self.dt already exists, using evolve_spring with a different one
                print("Warning: you are continuing with a different timestep")
                print("previous dt=%.4f"%self.dt)
                print("new dt=%.4f"%dt)
                self.dt = dt
        except AttributeError :
            # Define the attribute dt for the first time
            self.dt = dt    
        
        self.num_timesteps = num_timescales * num_timesteps_per_timescale
        
    def evolve_circuit(self, num_timescales, num_timesteps_per_timescale) :
        """Evolve the spring, populate the acceleration, velocity and position (a_vals, v_vals, x_vals)
        Parameters
        ----------
        num_periods : float
            number of periods to evolve over
        num_timesteps_per_period : int
            number of timesteps per period, defines the time resolution         
        """
        
        self._set_num_timesteps(num_timescales, num_timesteps_per_timescale)
        
        for timestep in np.arange(self.num_timesteps) :
            # Populate i_vals and q_vals with current current and charge on capacitor
            self.i_vals.append(self.i_now)
            self.q_vals.append(self.q_now)
            
            # Update "now" values: q_now, and i_now (i_now can be updated with a method)
            self.q_now += self.i_now * self.dt 
            # YOUR CODE HERE
            raise NotImplementedError()
```

<!-- grade_id: cell-310b2f6c96da39e5, points: 2, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
test_rcc = RCCircuit(1,1e-6,0,10)
assert(test_rcc.i_now == 10)
test_rcc.evolve_circuit(10,100)
assert(test_rcc.i_now < 1e-3)
```

**Discussion** (5 min): Food for thought -- you'll notice that the class has two lines that call the `self._set_i_now` method.  Can you identify where you might place a *single call* of `self._set_i_now` instead such that everything works?  Write out some takeaways from the discussion below (2 points).

<!-- grade_id: cell-38765ae3b01db270, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### Model the charging capacitor in an RC circuit (2 points -- in class)

In the cell below, 
- Define an instance of `RCCircuit` that corresponds to an RC circuit with the following initial conditions $R = 2\Omega$, $C = 20e-06$ Farad, $V_b=5$V, $Q_0=0$.  An appropriate name for this instance would be `rcc`.  
- Evolve the system over 5$\tau$ (i.e. $5 RC$), where $\tau$ is the relevant timescale with 1000 steps per $\tau$.

<!-- grade_id: cell-38a6f85045c409f4, points: 0, grade: false, solution: true -->
```python
# Define charging_rc_circuit here and evolve
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c3659f886275a80d, points: 2, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert('i_vals' in dir(rcc))
assert(rcc.dt == 4e-8)
```

### Calculate the power (2 points -- 5 min group work)
- Calculate the power output by the battery at each timestep, p_vals_battery
- Calculate the power dissipated by the resistor at each timestep, p_vals_resistor

Each of these depend on the attribute `i_vals`, which is a list of the current at each timestep.  You can either create an array from this list (recommended), or use a list comprehension if you want to practice this (and the power at each timestep would be a list).

<!-- grade_id: cell-4f5198300f36b8d5, points: 0, grade: false, solution: true -->
```python
# Calculate the power here
# YOUR CODE HERE
raise NotImplementedError()
p_vals_battery.sum()
p_vals_resistor.sum()
```

<!-- grade_id: cell-365e9814aac8959e, points: 2, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
if type(p_vals_battery) == list :
    assert(np.array(p_vals_battery).sum() - 12415 < 1)
else :
    assert(p_vals_battery.sum() - 12415 < 1)

if type(p_vals_resistor) == list :
    assert(np.array(p_vals_resistor).sum() - 6252 < 1)
else :
    assert(p_vals_resistor.sum() - 6252 < 1)
```

### Plot the charge, current, and power delivered by the battery (2 points -- 5 min group work)

Below is most of the code to make a 3 panel plot of charge, current, and battery power as a function of time.  
- Fill in the appropriate code to calculate the time in units of $\tau=RC$ (analogous to time in units of the period T from the preflight), `time_rc_units`.  
- Define yarrays with the appropriate lists/numpy arrays to make the described 3 panel plot.

<!-- grade_id: cell-452afc9496050c97, points: 0, grade: false, solution: true -->
```python
#  Plotting code here
fig, (ax1, ax2, ax3) = plt.subplots(3, figsize=(8,8))
fig.suptitle('Charging Capacitor', fontsize=18)

# Calculate the time in units of RC here (time_rc_units)
# YOUR CODE HERE
raise NotImplementedError()

axes = [ax1, ax2, ax3]
# Define yarrays
# YOUR CODE HERE
raise NotImplementedError()

ylabels = ['Charge [C]', 'Current [A]', 'Battery Power [W]']
colors = ['b', 'r', 'g']

for ax, yarray, ylabel, color in zip(axes, yarrays, ylabels, colors) :
    ax.plot(time_rc_units, yarray, color=color)
    ax.grid()
    ax.set_ylabel(ylabel, fontsize=14)

ax3.set_xlabel('Time in units of RC', fontsize=16)


```

<!-- grade_id: cell-2f88e78fa85f7d41, points: 2, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert(time_rc_units.size==rcc.num_timesteps)
assert(np.abs(time_rc_units[-1] - time_rc_units[-2] - rcc.dt/rcc.rc_constant) < 1e-5)
```

**Discussion of figure** (2 min -- in class):  Let us talk through what is going on in the figure above.  Use the cell below to write a "figure caption" talking through the figure. (1 points)

<!-- grade_id: cell-15f77e019dfb07cc, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Discussion** (2 min): Yet even more food for thought -- power is energy per unit time.  We can therefore calculate both the energy delivered by the battery to the circuit (converted from the chemical potential energy in the battery), and also the energy dissipated into heat by the resistor.  What does the energy delivered by the battery correspond to in the plot above (i.e. in the third panel)?

### Calculate the energy budget (4 points)

Calculate:
- The energy delivered by the battery, `energy_battery`, over the time the circuit evolved, 
- The energy stored on the capacitor, `energy_capacitor`, due to the collected charge by the end of the circuit's time evolution, 
- The energy dissipated by the resistor, `energy_resistor`, 
- The efficiency of the system (fraction of energy *not* lost to heat), `energy_efficiency`.

<!-- grade_id: cell-93ba908bf33bb13f, points: 0, grade: false, solution: true -->
```python
# Calculate energy budget here
# YOUR CODE HERE
raise NotImplementedError()
print("energy delivered by the battery: ", energy_battery)
print("energy stored in capacitor: ", energy_capacitor)
print("energy dissipated by resistor: ", energy_resistor)
print("fraction of energy not lost to heat: ", energy_efficiency)
```

<!-- grade_id: cell-09c55d494af41d4a, points: 4, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert(energy_efficiency < 1)
assert(abs(1 - energy_capacitor/energy_resistor) < 0.05 )
```

**Comment on the energy budget** (2 points -- likely take home) : In the cell below, comment on the energy budget you calculated.  In particular, note how energy divides between the capacitor and the resistor.  Comment on why the energy efficiency is not 1.  Does the energy budget change for different values of R and C if we allow the system to evolve for a long time?  (Note, you can quickly test this to numerically check what happens.)

<!-- grade_id: cell-d19cd3451f3891e7, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

##  Homework:  Discharging the capacitor (10 points total)

This is arguably a simpler version of what we did during the tutorial, but there are some things to keep track of when executing the task.  Note, you can generally use the same class, `RCCircuit`, to model a discharging capacitor starting from an already charged capacitor.  The main difference is that you will not have a battery present.  This is equivalent to having the voltage of the battery equal to 0V.  For homework, you will discharge the capacitor in the same modeled system as before.  See the description of the system below.

### Discharging a Capacitor

We will model the discharging of the capacitor we charged in the tutorial section.  See schematic of the system below.
<img src="http://www-personal.umich.edu/~gerdes/img/MC-PythonCircuit01.png" height="400" width="400">
Let's briefly review the relations.  In absence of a battery, we have,
\begin{equation}
\frac {dQ} {dt} = I = \frac{V_{cap}}{R}
\end{equation}
and
\begin{equation}
 V_{cap} = \frac{Q_{cap}}{C}
\end{equation}


Thus, we have our simple first order differential equation
\begin{equation}
\frac {dQ} {dt} = \frac{Q}{RC}
\end{equation}

Turning this into a discrete expression, we get 
\begin{equation}
\Delta Q = \frac Q {RC} \Delta t
\end{equation}

### Modify the state of the object, evolve the circuit (2 points)

In the cell below, let us first change the state of the object `charging_rcc` by removing the battery (i.e. set the attribute, `voltage_battery` to zero.  Time evolve this circuit for the same number of periods and number of timesteps as when we previously charged the capacitor (5 time periods, 1000 timesteps per time period).

<!-- grade_id: cell-6d7699b323657b0e, points: 0, grade: false, solution: true -->
```python
# Define discharging_rc_circuit here and evolve
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-233e92afbeccf9b2, points: 2, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert(len(rcc.i_vals) == 2 * 5 * 1000)
assert(np.abs(rcc.i_now) < 0.1 )
```

### Calculate the power (2 points)

In the cell below, calculate `power_resistor`, the power dissipated by the resistor at each timestep (note, this will be an array covering the timesteps during the charging up and discharging).

<!-- grade_id: cell-86fcaa9d84c4974f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-45a2585e25a7e60d, points: 2, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert(abs(power_resistor.sum()/p_vals_resistor.sum() - 2.) < 0.1)
```

### Plot the behavior of the circuit (2 points)

Make a three panel figure (three axes) showing the charge, current, and the power dissipated by the resistor as a function of time in units of RC (similar to the previous case).  Note, in absence of the battery, current has a negative sign, it has reversed directions.  (Hint: When making the figure, you have the charge and current from the timesteps when the battery was connected.  So, the total number of timesteps is **twice** that stored from the most recent `evolve_circuit` method call.  From your final plots, you should see both the charging up and discharging behavior over 10 full time periods.) 

<!-- grade_id: cell-e2285d7e16548d23, points: 2, grade: true, solution: true -->
```python
# Plot the behavior in the circuit here
fig, (ax1, ax2, ax3) = plt.subplots(3,1, figsize=(8,12))  

# First calculate the power dissipated by the resistor, the time_rc_units, then plot as before

# YOUR CODE HERE
raise NotImplementedError()
```

### Compare energy budget (2 points)

Calculate the total energy dissipated, `energy_dissipated` and the energy remaining in the capacitor, `energy_capacitor`.  

<!-- grade_id: cell-35a8fad8378a5d95, points: 0, grade: false, solution: true -->
```python
# Calculate energy budget here
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1732df4c5ac57268, points: 2, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert(energy_capacitor < 1e-7)
```

### Comment on your results (2 points)

In the cell below, briefly describe the figure (charge, current, power behavior) as though you were writing a caption.  Also, comment on the energy budget of the system.

<!-- grade_id: cell-a46b942410366031, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE