# Phys 260 Python Lab 1: Introduction to Python: Preflight

## Introduction

Each Python lab will start with a pre-flight exercise that walks through building some of the set up and tools ($\sim$ 30 min), followed by an in-class tutorial with time for Q+A (50 min) so you can walk through steps that will be necessary for the homework assignment you will submit ($\sim$ 3 hrs).  Each lab will contain starter code, similar to what you see below.  Please fill in the code to complete the pre-flight assignment in preparation for the in-class tutorial.  

Preflight ($\sim$30-60 min, 10 points) **Typically due: Thursdays 3pm EST**

*Preflight typically graded by Wednesday 5p EST*

In-class tutorial and Q+A ($\sim$ 50 min, 10 points) **Typically occurs: Fridays 12pm EST**

Homework assignment ($\sim$ 3-5 hrs, 30 points) **Typically due: Tuesdays 11:59pm EST**

*Homework typically graded by Thursday 5p*

When we grade your homework, we will not run your code. Once submitted, your notebook should have the outputs for all of your results.  Please do not include long outputs from debugging, beyond a few print statements and the requested visualizations (i.e. plots).

**Grading:** When we grade your notebook, we will convert the .ipynb file to an HTML file.  We will be using [nbgrader](https://nbgrader.readthedocs.io/en/stable/) to grade your notebooks.  

## Preflight summary
- Introduction to importing modules (numpy and matplotlib)
- Introduction to functions
- Introduction to plots
- Introduction to other python "objects" (lists, dictionaries, strings)
- Introduction to loops

## Tutorial summary
- Plotting the one dimensional electric field due to a point charge
- Comparing a "model" to the data
- Plotting the electric field in 3 dimensions (the meshgrid method for generating points and 3d plotting tools)

## Homework summary
- Plot the electric field due to a dipole in 3 dimensions
- Plot the same electric field along 2d planes

## Importing Modules

**Run (execute) the code in the following cell using the combined keys, shift+enter, after clicking in the cell.** You can also click the "Run" button.  

This code will import the modules needed for this pre-flight.  Modules contain a series of tools.  The [numpy](https://numpy.org/) module is actually an entire package comprised of tools for NUMerical PYthon.  The [matplotlib](https://matplotlib.org/) module is also a package, but this one has a collection of tools in a MAThematical PLOTting LIBrary. We will use the pyplot module from this library. You can think of modules as a collection of already written code (tools) that you get to use.  Libraries are collections of modules.



```python
import numpy 
from matplotlib import pyplot
```

We can now use access tools in either the numpy library or the matplotlib library as follows, **using the imported name as a prefix**. Next, execute the code in the follow cell. 

```python
# Create an array of sequential numbers
example_array = numpy.arange(1,10)
```

```example_array``` is a variable pointing to a *python object* that is a numpy n-dimensional array.  

In executing the next cell, we (1) print the type of the python object, (2) print the contents of the python object, (3) *output* the python object.  

Note, lines preceded with a hash ```#``` denote *comments* in the code, and are not executed lines.

```python
# Print the type of the object
print('This object is of type: ', type(example_array))

# Print the contents of the object
print('The contents of this object are: ', example_array)

# Let us output the contents of example_array
example_array
```

Output:
```
This object is of type:  <class 'numpy.ndarray'>
The contents of this object are:  [1 2 3 4 5 6 7 8 9]
array([1, 2, 3, 4, 5, 6, 7, 8, 9])
```

---
## Module import with alias exercise (1 point)

In the next cell, try importing numpy and pyplot using the standard aliases.  This means executing code of the form:
```
import [module] as [another_name]
```

The standard alias for numpy is np, and the standard alias for pyplot is plt.  Note, after filling out your solution, remove the line:

```raise NotImplementedError()```

<!-- grade_id: aliases, points: 0, grade: false, solution: true -->
```python
# Import numpy and pyplot with aliases here
### BEGIN SOLUTION
import numpy as np
from matplotlib import pyplot as plt
### END SOLUTION

```

<!-- grade_id: Module import, points: 1, grade: true, solution: false -->
```python
"""Execute this cell to check that pyplot has been imported with the appropriate aliases"""
assert(id(pyplot)==id(plt))
### BEGIN HIDDEN TESTS
# Checks that numpy has been imported with appropriate alias
assert(id(numpy)==id(np))
### END HIDDEN TESTS
```

---
## Module alias test (1 point)

Given the alias for the module name you just imported, you will now want to try this out.  Write the following code in the next cell and execute the cell in order to define the variable ```xarray```
```
xarray = np.arange(1,10)
```

<!-- grade_id: xarray, points: 0, grade: false, solution: true -->
```python
# Define xarray here
### BEGIN SOLUTION
xarray = np.arange(1,10)
### END SOLUTION
```

<!-- grade_id: Module alias test, points: 1, grade: true, solution: false -->
```python
"""Execute this cell to check that xarray[0] returns the correct output"""
assert(xarray[0]==1)
### BEGIN HIDDEN TESTS
assert((xarray==np.arange(1,10)).all)
### END HIDDEN TESTS
```

## Introduction to functions/python methods

You can think of functions/methods as python objects that *do* something (think of these as a verb).  As *doing* objects, we (the user) executes them.  The arguments, or keyword arguments of a function/method say *how* to do something (think of these as an adverb).  

In the cell above, you executed the **arange function from numpy** (aliased to np), and provided the **arguments 1 and 10**.  The instructions underlying this function are: **return** an n-dimensional numpy array starting at 1 and ending before 10.  There is a **default keyword argument, step,** and it defaults to 1.  

```
np.arange(1, 10, step=0.5)
```
This also works:
```
np.arange(1,10,.5)
```

Try outputting one of the two lines above that use the step keyword argument in the cell below.

```python
#  Execute this cell, which uses the step keyword argument in the np.arange method, outputting the returned object
np.arange(1,10, step=0.5)
```

In the next cell, we (1) print the type of the python object ```np.arange``` un-executed (you should see that this is a built in function/method), (2) print the contents of the python object returned by this function 

```python
print(type(np.arange))
print(np.arange(1,10,.5))
```

---
## Explain a keyword argument (1 point)

In words in the markdown cell (note: it is not a code cell) below, describe in words what the step key word argument does.  e.g. "The keyword step in the arange function serves to.... "

<!-- grade_id: Step Explanation, points: 1, grade: true, solution: true -->
The keyword "step" in the arange function indicates the difference between successive numbers in the array.

### Anatomy of a function

Below we build our own function.  When you build your own functions, it is **good practice to give it a name (starting with a verb) that clearly describes what it does**.  The ```raise_to_power``` function takes in an argument base, and a keyword argument power with default value 2.  This returns the base raised to the power.

You'll notice that is text enclosed in three quotes on either side.  This is known as a docstring (documentation string).  Don't forget to execute the code below so the function is actually defined.

```python
def raise_to_power( base, power=2) :
    """Raise the base to the power and return the results."""

    return base**power
```

Here, we test our function with assertion statements.  The ```assert``` function takes in a boolean argument, something that returns True or False.  Common boolean arguments contain comparison operations, such as ```==, !=, <, >, <=, >=```

In general, tests should check that functions do what you expect them to do.

```python
assert(raise_to_power(2)==4)
assert(raise_to_power(2,3)==8)
```

---
## Using a function (1 point)
One exciting thing to note is that entire numpy arrays can be raised to a power!  In the assertion statements above, we tried inputting 2 for the base. 

In the cell below, input xarray as the first argument to your ```raise_to_power``` function, and leave the keyword argument to its default.  The result should be xarray squared one by one.  Set a new variable ```yarray``` equal to this, and print the contents of ```yarray```

<!-- grade_id: function-use, points: 0, grade: false, solution: true -->
```python
# Set yarray equal to the square of xarray here, e.g. yarray = ...
### BEGIN SOLUTION
yarray = raise_to_power(xarray)    
### END SOLUTION
```

<!-- grade_id: Using a function, points: 1, grade: true, solution: false -->
```python
"""Execute this cell to check that yarray is correctly defined"""
assert(yarray[0] == raise_to_power(xarray[0]))
### BEGIN HIDDEN TESTS
assert((yarray == raise_to_power(xarray)).all)
### END HIDDEN TESTS
```

## Quick plot

Next, we will plot yarray vs. xarray with a quick tool from pyplot (aliased to plt) called plot.  This creates a plot of a single axis.  After we plot the line, we adjust the plot to add x and y labels, specifying how we add them (e.g. with the fontsize key word argument).   

Look at the documentation for [pyplot.plot](https://matplotlib.org/3.3.0/api/_as_gen/matplotlib.pyplot.plot.html), and modify the code below to change the color of the line.  Hint:  You'll need to add a keyword argument.

```python
# Plot the arrays against one another with the simple plot function from pyplot
plt.plot(xarray, yarray)
plt.xlabel('x array values', fontsize='xx-large')
plt.ylabel('$x^2$', fontsize='xx-large')
```

## Using a method (1 point)

Methods are pretty similar to functions.  The main difference is that a method is inherently tied to a python object.  Let's take the example of the n-dimensional numpy array object.  We have two of these, xarray and yarray.   One method of the n-dimensional numpy array object is ```mean```, which takes the mean of that array. 
```
yarray.mean()
```
Note, here we see that the above line executes the ```mean``` method of an ```ndarray``` object whose name is yarray.  Note, we did not provide any arguments to the method; no arguments were needed.  We can also use the numpy function mean, ```np.mean``` and provide the numpy array as an argument, e.g. 
```
np.mean(yarray)
```
To calculate the median of an array, there is no method of the n-dimensional numpy array object for the median, but there is a function in the numpy library called median.   

In the following cell, define two more variables ```ymean``` and ```ymedian``` that correspond to the mean and median of ```yarray```.  Hint, you will need to use the numpy function median.

<!-- grade_id: mean-median, points: 0, grade: false, solution: true -->
```python
# Define ymean and ymedian here, e.g. ymean = ..... 
### BEGIN SOLUTION
ymean = yarray.mean()
ymedian = np.median(yarray)
### END SOLUTION
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-4-186f3847e1c4> in <module>
      1 # Define ymean and ymedian here, e.g. ymean = .....
      2 ### BEGIN SOLUTION
----> 3 ymean = yarray.mean()
      4 ymedian = np.median(yarray)
      5 ### END SOLUTION

NameError: name 'yarray' is not defined
```

<!-- grade_id: Using a Method, points: 1, grade: true, solution: false -->
```python
"""Example check that the variables ymean are correctly set"""
assert(ymean==yarray.mean())
### BEGIN HIDDEN TESTS
assert(ymedian==np.median(yarray))
### END HIDDEN TESTS
```

---
## Plotting exercise (2 points)

In the next cell, plot the same curve as before (yarray vs. xarray), and look up how to plot horizontal lines.  You will plot a horizontal line corresponding to the ymean and ymedian (*Hint: look up "Plot horizontal line matplotlib"*).  Add a legend (look this up as well, *Hint: look up "Plot legend with matplotlib*) that shows which line is the mean and which line is the median.

<!-- grade_id: Plotting exercise, points: 2, grade: true, solution: true -->
```python
#  Create your plot here
### BEGIN SOLUTION
plt.plot(xarray, yarray)
plt.hlines(ymean, xmin=xarray.min(), xmax=xarray.max(), label='mean', color='r')
plt.hlines(ymedian, xmin=xarray.min(), xmax=xarray.max(), label='median', color='k')
plt.xlabel('x array values', fontsize='xx-large')
plt.ylabel('$x^2$', fontsize='xx-large')
plt.legend()
### END SOLUTION
```

## Other python objects

Here, we will briefly look at some other standard python objects that may be useful as you learn to code in python.

- [Strings](https://docs.python.org/2/library/string.html): Strings are collections of characters, enclosed in quotations.  We often use these to print information or label things in visualizations.  Strings are mutable (see below).
```
example_string = "hello world"
```

- [Lists](https://docs.python.org/3/tutorial/datastructures.html): lists are an ordered collection of python objects, enclosed in square brackets ```[]```, where each object is delimeted (separated) by a comma.  The following list contains two integers, an n-dimensional numpy array, a string of characters, and a list with two items.  Lists are [*mutable*](https://medium.com/@meghamohan/mutable-and-immutable-side-of-python-c2145cf72747) objects, so they can be changed after first created.  
```
example_list = [0, 1, xarray, 'hello world', [100,200]]
```

- [Tuples](https://docs.python.org/3/tutorial/datastructures.html): tuples are an ordered collection of python objects, enclosed in round braces ```()```, where each object is delimeted by a comma.  These are *immutable*.
```
example_tuple = (0, 1, xarray, 'hello world', [100, 200])
```

- [Dictionaries](https://docs.python.org/3/tutorial/datastructures.html): dictionaries are key, value paired collections of python objects, enclosed in squiggly brackets ```{}```, where each pair is delimeted by a comma, and the pairing is indicated with a colon.  Below are two ways to explicitly define a dictionary.  Our example dictionary below happens to have keys and values that are all strings.  Note, dictionaries, like lists, are mutable.
```
example_dictionary = {'feline': 'cat-like', 'canine': 'dog-like'}
```
or using the built-in dict function
```
example_dictionary = dict([('feline', 'cat-like'), ('canine':'dog-like')])
```

---
## Create python objects (1 point)

In the cell below, use the above code to define ```example_string, example_list, example_tuple, and example_dictionary```.  

<!-- grade_id: example-objects, points: 0, grade: false, solution: true -->
```python
# Define example objects here
### BEGIN SOLUTION
example_string = "hello world"
example_list = [0, 1, xarray, 'hello world', [100,200]]
example_tuple = (0, 1, xarray, 'hello world', [100, 200])
example_dictionary = {'feline': 'cat-like', 'canine': 'dog-like'}
### END SOLUTION
```

<!-- grade_id: Create python objects, points: 1, grade: true, solution: false -->
```python
"""Check that example objects have correct output"""
assert(example_string == "hello world")
assert(example_list == [0, 1, xarray, 'hello world', [100,200]])
### BEGIN HIDDEN TESTS
assert(example_tuple == (0, 1, xarray, 'hello world', [100, 200]))
assert(example_dictionary == {'feline': 'cat-like', 'canine': 'dog-like'})
### END HIDDEN TESTS
```

---
## Try methods of python objects (1 point)

You can create a version of example string with,
```capitalized_example_string = example_string.capitalize()```
Execute this below, and print the contents of both capitalized_example_string and example_string.  You'll notice that the original example_string did not change.  The ```str.capitalize``` method does *not* change the string in-place, but instead copies the string and returns a changed copy.

Create a new variable ```appended_example_list``` by setting it equal to the return value of ```example_list.append(capitalized_example_string)```.  Print the contents of both ```appended_example_list``` and ```example_list```. You'll notice that the original example_list has changed.  What is the printed output of ```appended_example_list```, or rather the output of the append method of the list object (i.e. ```list.append```)?  What we are seeing here is that the ```list.append``` method *does* change the list in place.  It changes the original list.  But, it returns None (nothing, the default return value of any function/method if no return value is specified).

<!-- grade_id: modify-objects, points: 0, grade: false, solution: true -->
```python
# Define capitalized_example_string and appended_example_list here, e.g. capitalized_example_string = ...
### BEGIN SOLUTION
capitalized_example_string = example_string.capitalize()
print(capitalized_example_string)
print(example_string)

appended_example_list = example_list.append(capitalized_example_string)
print(example_list)
print(appended_example_list)
### END SOLUTION
```

<!-- grade_id: Try methods of python objects, points: 1, grade: true, solution: false -->
```python
"""Check that modified objects have correct output"""
assert(capitalized_example_string == example_string.capitalize())
### BEGIN HIDDEN TESTS
assert(appended_example_list == None)
### END HIDDEN TESTS
```

## Test mutable vs. immutable things (1 point)

We will illustrate mutability/immutability with dictionaries and tuples, respectively.  As an example, let us first consider lists, which are mutable.  We can access an individual item of a list using that item's index, e.g.

```python
print(example_list[1])
```

The above printed the integer 1.  We can re-assign the value in the 1st index of this list and print its contents with,

```python
example_list[1] = "first index of list!"
print(example_list[1])
```

The above should have printed the string, ```"first index of list!"```.  Let us see how to do similar things for dictionaries and tuples.


Note, you can access individual elements of a dictionary using their keys, and individual elements of a tuple using the item index number (similar to lists). For example, 


```python
print(example_dictionary['feline'])
print(example_tuple[1])
```

The above should have printed ```"cat-like"``` in the first line and the integer 1 in the second line, similar to the case with the list.  Try reassigning (i.e. changing) the value for ```example_dictionary['feline']``` so it instead outputs "lion-like".  Then, try to re-assign the value of ```example_tuple[1]``` so it instead outputs "first index of tuple!".  

<!-- grade_id: modify-dict-item, points: 0, grade: false, solution: true -->
```python
# Re-assign the value for example_dictionary['feline'] here, then print the output.
### BEGIN SOLUTION
example_dictionary['feline'] = "lion-like"
print(example_dictionary['feline'])
### END SOLUTION

```

<!-- grade_id: Test mutable vs. immutable things, points: 1, grade: true, solution: false -->
```python
"""Check that modified objects have correct output"""
assert(example_dictionary['feline'] == "lion-like")
```

```python
# Try to re-assign the value for example_tuple[1] here
```

Note, you should have gotten a ```TypeError``` above.  Tuples cannot be changed after you have created them.  They are immutable.

## Introduction to for loops - and quick note on when *not* to use them

One key tool in programming that can save extra unnecessary code is the [for loop](https://www.w3schools.com/python/python_for_loops.asp).  A for loop simply iterates a set of instructions over every item in an iterable, where an iterable is any object in python that can be accessed one item at a time.  Strings, lists, dictionaries, and tuples are all iterables.  A loop will look something like,

```
for list_item in example_list :
    # Prints the type of python object for each object in example_list
    print(type(list_item))
```
or, we can simultaneously iterate over the key, value pairs in a dictionary: 
```
for key, value in example_dictionary.items() :
    # Re-assigns the value to each dictionary item to a version that is capitalized
    example_dictionary[key] = value.capitalize()
```


**Note:** Numpy n-dimensional arrays are also iterable.  *But*, if you are dealing with arrays, it is often the case that you should be able to vectorize whatever set of instructions you with to complete.  We show an example of when to use them and when not to use them below

### Loop over things to plot
One reasonable application of a loop is if you are plotting similar things with minor changes.  Below, we loop over ```different_powers```, which is a list of integers, and plot that power as a function of the contents of ```xarray```.  Note two new things to learn from the code below.  
- First, we created a string for the legend, ```legend_label```, that included a string representation of the ```different_power``` we are iterating on.  This is a useful syntax for printing
- Second, we are using the method of pyplot (alias plt) called yscale.  This sets the scale, and we have used the argument 'log' to specify that we want the yscale to be log-spaced.

```python
#  When to use a loop
different_powers = [2, 3, 4, 5]
for different_power in different_powers :
    legend_label = 'power=%i'%different_power
    plt.plot(xarray, raise_to_power(xarray, power=different_power), label=legend_label)

    
plt.yscale('log')  #  Note, when dealing with a large dynamic range of numbers, plot things on a log scale
plt.legend()
```

### Loop over elements in a numpy array (a no-no!)

Below, we loop over elements in ```xarray``` to raise each individual element to a power, then use this to populate the yvalues (using list appending), then converting that back to a numpy array.  This is a no-no.  Do. not. do. this.  As you saw in previous lines of code, you can simply do,
```
plt.plot(xarray, raise_to_power(xarray))
```
This is because we can do the raising to a power step in a *vectorized* way (all at once). 

```python
yvalues = []
for xvalue in xarray :
    yvalues.append(raise_to_power(xvalue))
    
plt.plot(xarray, yvalues)
```