# Phys 260 Python Lab 5: Python Classes -- Preflight into a bit of history

# Introduction

Each Python lab will start with a pre-flight exercise that walks through building some of the set up and tools ($\sim$ 30 min), followed by an in-class tutorial with time for Q+A (50 min) so you can walk through steps that will be necessary for the homework assignment you will submit ($\sim$ 3 hrs).  Each lab will contain starter code, similar to what you see below.  Please fill in the code to complete the pre-flight assignment in preparation for the in-class tutorial.  

Preflight ($\sim$30-60 min, 10 points) **Typically due: Wednesdays 3pm EST**

*Preflight typically graded by Wednesday 5p EST -- see your feedback in html (launch a browser)*

In-class tutorial and Q+A ($\sim$ 50 min, 10 points) **Typically occurs: Thursdays 9am EST**

Homework assignment ($\sim$ 3-5 hrs, 30 points) **Typically due: Mondays 9a EST**

*Homework typically graded by Thursday 5p -- see your feedback in html (launch a browser)*

When we grade your homework, we will not run your code. Once submitted, your notebook should have the outputs for all of your results.  Please do not include long outputs from debugging, beyond a few print statements and the requested visualizations (i.e. plots).

**Grading:** When we grade your notebook, we will convert the .ipynb file to an HTML file.  We will be using [nbgrader](https://nbgrader.readthedocs.io/en/stable/) to grade your notebooks.  

## Preflight summary
- Reminder of the familiar `ndarray` class
- Example case with a `Person`

```python
# Import relevant modules
import numpy as np
from matplotlib import pyplot as plt
```

## Quick reminder of your use of the ndarray object - a python class

A python class bundles data and functionality (things and actions). One example of a built-in python class you've been using is the ndarray class.  See this in action below, where we print the `type` of the object `test_array`.  Note:  You can think of a class as the blueprint/template for an object.

```python
test_array = np.arange(10)
print(type(test_array))
```

Output:
```
<class 'numpy.ndarray'>
```

### Familiar attributes and methods

In the cell above, we have defined an instance of the `ndarray` class, with name `test_array`.  Each class instance has its own attributes attached to it. Class instances can also have methods that perform actions on class attributes.

Name 2 attributes and 2 methods of `ndarray` objects.  

**Answer**: 
- Attributes - size, shape
- Methods - min(), mean()

### Define a method in a class (2 points)

We start with a canonical example of a class, a `Person` with attributes name and age.  Many intro to python classes use examples like people, cars, cities.  We'll use people.

Take a look at the `class` definition below.   Each `Person` is initialized with attributes `self.name` and `self.age`.  The initial state of the class is defined in the `__init__` method.  So, the initial state has these two attributes.  There are also three methods in this class, `self.introduce_person`, `self.inquire_contribution`, and `self.add_relationship`.  


The attributes, indicate keep track of the `state` of a given instance of this class.  The methods perform actions on the class attributes, and can change the `state`. 

Define another method of the `Person` class to `add_relative`, where this method defines the attribute `self.relative` and takes in `relative_to_add` as an argument.

<!-- grade_id: cell-e694ba7b42b338d3, points: 0, grade: false, solution: true -->
```python
class Person :
    def __init__(self, name, age) :
        self.name = name
        self.age = age
        not_an_attribute = 'purple monkey dishwasher'
    
    def introduce_person(self) :
        """Introduces self"""
        print("Hi, my name is %s"%self.name +" I am %d"%self.age)
    
    def inquire_contribution(self) :
        """Indicate the historical physics/math/chemistry contribution"""
        try :
            print("I %s"%self.contribution)
        except AttributeError :
            print("No contribution yet added")
    
    def add_contribution(self, contribution_to_add) :
        """Add contribution this person made"""
        self.contribution = contribution_to_add
    
    ### BEGIN SOLUTION
    def add_relative(self, relative_to_add) :
        """Add relative"""
        self.relative = relative_to_add
    ### END SOLUTION
```

<!-- grade_id: cell-48a9bf6f430602b5, points: 2, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert("add_relative" in dir(Person))
```

### List class attributes and methods (2 points)

- List the attributes of `Person` that exist when you instantiate the class.
- List the methods of `Person`.
- List the attributes of `Person` that will exist after executing other methods.

<!-- grade_id: cell-6f601fc8eafa7780, points: 2, grade: true, solution: true -->
**Answer:**
- Person.name and Person.age are attributes
- Person.introduce_person, Person.inquire_contribution, Person.add_contribution, and Person.add_relative are all methods.
- Person.contribution and Person.relative are attributes that will exist after executing other methods.

## Class instances

Below, we define four instances of the `Person` class.  Each one of these people are historical physicists, and we use their final age in the `self.age` attribute.  I picked out two pairs of physicists who are related to one another and share a last name - Richard and Joan Feynman are siblings, and Marie and Pierre Curie are spouses.  The former of each are the more historically famous of each pair, with Richard Feynman and Marie Curie being media popularized Nobel Laureates.  

Extra fun fact, in case you find yourself doing trivia night on zoom: Marie Curie received **two** [Nobel prizes](https://en.wikipedia.org/wiki/Marie_Curie#Nobel_Prizes) - one in physics (1903) and one in chemistry (1911).   Pierre Curie also received the 1903 Nobel prize, but Marie Curie's double award likely boosted her historical fame.

```python
richard = Person('Richard Feynman', 69)
joan = Person('Joan Feynman', 93)
marie = Person('Marie Curie', 66)
pierre = Person('Pierre Curie', 46)
print(type(joan))
```

Output:
```
<class '__main__.Person'>
```

You'll notice that in the cell above, the `type` of any of the instances is class `'__main__.Person'`.  Compare this with the `np.ndarray` class when we did the same thing for `test_array`.  The `__main__` module is the module we are currently "in", i.e. this jupyter notebook.  We defined the `Person` class here.  The `ndarray` class is defined in the `numpy` module.

## Use a method in your object

In the cell below, we execute one of the methods for the instance `richard`.  We then start a `for` loop to do the same for all instances of the `Person` class we have defined.  Execute the same method in the `for` loop so everyone gets introduced.

<!-- grade_id: cell-f40363f5fc9e1711, points: 1, grade: true, solution: true -->
```python
richard.introduce_person() 

for physicist in [richard, joan, marie, pierre] :
    ### BEGIN SOLUTION
    physicist.introduce_person()
    ### END SOLUTION
```

Output:
```
Hi, my name is Richard Feynman I am 69
Hi, my name is Richard Feynman I am 69
Hi, my name is Joan Feynman I am 93
Hi, my name is Marie Curie I am 66
Hi, my name is Pierre Curie I am 46
```

In the cell below, we execute another method, `Person.inquire_contribution`.  You'll notice that there is no attribute `Person.contribution` yet. 

```python
richard.inquire_contribution()
```

Output:
```
No contribution yet added
```

### Modify an object's state:  ndarray case (1 point)

You have modified an object's state before when changing the value in an array element.  In the cell below, modify the last element of `test_array` to be 20 instead of 9 and print `test_array`.   

<!-- grade_id: cell-9197a1cd1ae6ecc5, points: 0, grade: false, solution: true -->
```python
# Modify the last element of test_array to be 20
### BEGIN SOLUTION
test_array[-1] = 20
print(test_array)
### END SOLUTION
```

Output:
```
[ 0  1  2  3  4  5  6  7  8 20]
```

<!-- grade_id: cell-55252a3d2043e1f3, points: 0.5, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert(test_array[-1] == 20)
```

## Modify an object's state:  Using the explicit method

Note, setting an element uses an under-the-hood method of ndarray:  `ndarray.__setitem__` ([see docs](https://numpy.org/doc/stable/reference/generated/numpy.ndarray.__setitem__.html)).  You'll notice that this particular method, as opposed to something like `ndarray.mean`, is enclosed by double underscores.  These are otherwise known as dunder (double under).  The dunder typically indicates that the method is an **under-the-hood method not to be explicitly used**.  Nonetheless, you'll explicitly use this in the cell below so you can see how this is equivalent to what you did above.  On a related note, the `__init__` is another such method, which happens under-the-hood when you instantiate a class (i.e. when you defined `richard`, `joan`, `marie`, and `pierre`.  

In the cell below, use the `__setitem__` method to change the first element of `test_array` to be 10 instead of 0.  Print `test_array` to see the change.  The method takes in two arguments - the first is the index of the element to be set, and the second is the value to set the element to.

<!-- grade_id: cell-741d988beee241cb, points: 0, grade: false, solution: true -->
```python
# Set the 0th element to 10
### BEGIN SOLUTION
test_array.__setitem__(0,10)
print(test_array)
### END SOLUTION
```

Output:
```
[10  1  2  3  4  5  6  7  8 20]
```

<!-- grade_id: cell-056335350584fa85, points: 0.5, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert((test_array == np.array([10,1,2,3,4,5,6,7,8,20])).all())
```

## Modify an object's state: Person case (1 point)

If you try printing out or outputting the attribute `Person.contribution`, you'll get an `AttributeError`.  Try this in the cell below by just removing the `NotImplementedError` and executing the cell.

Now, use the `add_contribution` module.  You can input as an argument, ["Fundamental work in quantum electrodynamics."](https://www.nobelprize.org/prizes/physics/1965/feynman/facts/)  As was mentioned during lecture, quantum electrodynamics is the treatment of E&M as we approach smaller and smaller scales, forming the theory of the interaction of quantized electromagnetic fields.   

Once you execute the cell below, you'll see the contribution properly printed.  

<!-- grade_id: cell-ec7d39485f1b0aa9, points: 0, grade: false, solution: true -->
```python
# An attribute
### BEGIN SOLUTION
richard.add_contribution("laid foundations for quantum electrodynamics")
### END SOLUTION
print(richard.contribution)
```

Output:
```
laid foundations for quantum electrodynamics
```

<!-- grade_id: cell-fdb9565276766582, points: 1, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert('contribution' in dir(richard))
```

### Aside of control flow
Recall, when we last inquired the contribution of a physicist, we got "No contribution yet added" printed to screen instead of an `AttributeError`.  That is because we had a `try` and `except` statement.  This is analogous to an `if` and `else`.  
```
if test_array[0] == 10 :
    print("I changed the first element")
else :
    print("I did not change the first element")
```

These are further examples of [control flow](https://www.oreilly.com/library/view/python-in-a/0596001886/ch04s09.html) in python.  You've already seen examples of `for` loops and `while` loops.  The `try` and `except` statement allows you to keep the program from throwing an error and exiting (if you predict the error, and in this case we are predicting an `AssertionError`) and to do something else when the predicted error is thrown (in this case, print 'No contribution yet added').   See below.

```python
richard.inquire_contribution()
joan.inquire_contribution()
```

Output:
```
I laid foundations for quantum electrodynamics
No contribution yet added
```

### Use other methods (2 points)

In the cell below, do the following:

- Add relatives to all instances of `Person` (i.e. all objects corresponding to the four physicists), analogous to how you added Richard Feynman's contribution above. Joan would have "Richard: Sibling" (and vice versa), Pierre would have "Marie: Spouse" (and vice versa) 
- Add Marie Curie and Pierre Curie's contribution that led to the Nobel Prize awarded to both, ["established an understanding of radioactivity"](https://www.nobelprize.org/prizes/physics/1965/feynman/facts/).  Unfortunately, these experiments likely led to Marie Curie's early death due to aplastic anaemia linked to long-term exposure to radiation.  Pierre Curie died an even earlier death due to a road accident in 1906.

<!-- grade_id: cell-44f857b4b53a445f, points: 0, grade: false, solution: true -->
```python
# Add relationships and one of Marie Curie's contributions here
### BEGIN SOLUTION
richard.add_relative("Joan: Sibling")
joan.add_relative("Richard: Sibling")
marie.add_relative("Pierre: Spouse")
pierre.add_relative("Marie: Spouse")
marie.add_contribution("established an understanding of radioactivity")
pierre.add_contribution("established an understanding of radioactivity")
### END SOLUTION
```

<!-- grade_id: cell-f0a948d10b1f3769, points: 1, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
for physicist in [richard, joan, marie, pierre] :
    assert('relative' in dir(physicist))
assert('contribution' in dir(marie))
assert('contribution' in dir(pierre))
marie.inquire_contribution()
pierre.inquire_contribution()
### BEGIN HIDDEN TESTS
marie_contribution_1 = marie.contribution
### END HIDDEN TESTS
```

Output:
```
I established an understanding of radioactivity
I established an understanding of radioactivity
```

### Change the state again (1 point)
In the cell below, we change the state of the instance, `marie`, yet again.  Add Marie Curie's contribution that led to the Nobel Prize in 1911, having "produced radium as a pure metal, proving the new element's existence".  This particular scientific advancement is the basis on which radioactive compounds are used to treat tumors in modern-day medicine.

<!-- grade_id: cell-c7ee799877ff736c, points: 0, grade: false, solution: true -->
```python
# Add Marie Curie's contribution for the second Nobel prize here
### BEGIN SOLUTION
marie.add_contribution("produced radium as a pure metal, proving the new element's existence.")
### END SOLUTION
```

<!-- grade_id: cell-21f1f038c340551b, points: 1, grade: true, solution: false -->
```python
"""Execute to check you're on teh right track"""
assert('contribution' in dir(marie))
marie.inquire_contribution()
### BEGIN HIDDEN TESTS
assert(marie_contribution_1 != marie.contribution)
### END HIDDEN TESTS
```

Output:
```
I produced radium as a pure metal, proving the new element's existence.
```

You'll notice that Marie Curie's first Nobel awarded contribution was completely overwritten by the second, not added to the text.  Food for thought - how might you modify the original definition of the class `Person` such that you actually add contributions instead of overwriting?

### Final thoughts

Joan Feynman was awarded a NASA Exceptional Achievement medal for her contributions in solar wind particles and fields.  As an example contribution, Joan Feynman "developed an understanding of the origin of auroras."  This is probably one of the coolest things to have studied IMO.  (See [really cool NASA infographic](https://spaceplace.nasa.gov/aurora/en/) for the general public - something definitely shareable with friends and family not explicitly studying physics).  Auroras are a visually appealing example of charged particles in a magnetic field, which you'll learn more about later.

```python
joan.add_contribution("developed an understanding of the origin of auroras")
joan.inquire_contribution()
```

Output:
```
I developed an understanding of the origin of auroras
```

[Joan Feynman](https://en.wikipedia.org/wiki/Joan_Feynman) passed away this last summer at the age of 93.  RIP.