# Phys 260: Python assignment header

### (1) Fill out the cell below.  
The cell below is a **code cell**.  Fill out your University of Michigan uniqname, then your name, and collaborators in the cell below **inside the quotes**.  

**Do not delete the quotes.**  We will use this information to organize your assignments.  To edit and execute cells, double click inside the cell, type, and press \<shift\>+\<enter\> to execute.

```python
UNIQNAME = ""
NAME = ""
COLLABORATORS = ""
```

### (2) Check your python version.  
**Execute the cell below** (double click in the cell and press \<shift\>+\<enter\>, or click in the cell and press the Run button) to check that you are using a version of python that is compatible with the tool we are using to grade your assignments.  If your ```IPython``` version is too old, we will *not* be able to grade your assignments.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

### (3) Do your best to answer all questions in the assignment.  
To answer questions, **replace** anything that says either
- "YOUR ANSWER HERE" 
- 
```
YOUR CODE HERE
raise NotImplementedError
``` 

with your answer/code.  Cells with either of the two bullet points above are cells of the notebook that will be graded.

**To edit markdown** cells (e.g. this one),  *double click in the cell to type*.  Press \<shift\>+\<enter\> to execute the cell.  Try editing the text below to replace the with your information (e.g. Camille Avestruz, cavestru):  

[first name] [last name], uniqname


### (4) Make sure your notebook runs sequentially.
After you complete this assignment, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

```python

```

# Phys 260 Python Lab 2: Numpy tools for vectorized actions

# Introduction -- Reminder

Each Python lab will start with a pre-flight exercise that walks through building some of the set up and tools ($\sim$ 30 min), followed by an in-class tutorial with time for Q+A (50 min) so you can walk through steps that will be necessary for the homework assignment you will submit ($\sim$ 3 hrs).  Each lab will contain starter code, similar to what you see below.  Please fill in the code to complete the pre-flight assignment in preparation for the in-class tutorial.  

Preflight ($\sim$30-60 min, 10 points) **Typically due: Thursdays 3pm EST**

*Preflight typically graded by Fridays 5p EST*

In-class tutorial and Q+A ($\sim$ 50 min, 10 points) **Typically occurs: Fridays EST**

Homework assignment ($\sim$ 3-5 hrs, 30 points) **Typically due: Tuesdays 12pm EST**  


When we grade your homework, we will not run your code. Once submitted, your notebook should have the outputs for all of your results.  Please do not include long outputs from debugging, beyond a few print statements and the requested visualimzations (i.e. plots).

**Grading:** Your notebook will be graded using [nbgrader](https://nbgrader.readthedocs.io/en/stable/), implemented using Gradescope's autograder function .  **Note:** Execute the cell below (click in the cell and press shift+enter, or click in the cell and press the Run button) to check that you are using a version of python that is compatible with the tool we are using to grade your assignments.  If your ```IPython``` version is too old, we will *not* be able to grade your assignments.


## Preflight summary
- Brief review of creating arrays
- Creating an (n,m,...) shaped array and examining the contents (2 points)
- Example of array operations using explicit power notation and a logarithm. (2 points)
- Calculate the electric field in a vectorized fashion (loop vs. one-liner!). (3 points)
- Calculate efield magnitude with different dimensions (1 point)
- Quiver plot (1 points) 

```python
# Import relevant modules
import numpy as np
from matplotlib import pyplot as plt
```

## Brief review of creating arrays

Numpy arrays are simply a grid of values, which are 0-indexed.  There are a number of ways to create a numpy array.  Below, we show: 
- Creating a numpy array from a list
- Creating a numpy array using `arange`
- Creating a numpy array using `linspace`

These are all shape (n,), essentially 1-d arrays.

```python
# Manually input a list of values and enclose in np.array()
test_list = [0.5, 1, 0, 20]
array_from_list = np.array(test_list)
print("array from list: ",array_from_list)
print("shape: ", array_from_list.shape)

# Use a numpy function to generate a certain kind of array
array_from_arange = np.arange(1, 10, 0.5)
print("array from arange: ", array_from_arange)
print("shape: ", array_from_arange.shape)

array_from_linspace = np.linspace(1,10,20)
print("array from linspace: ", array_from_linspace)
print("shape: ", array_from_linspace.shape)
```

## Creating an (n,m,...) shaped array and examining the contents (2 points)
With the numpy function `meshgrid`, we can also create shape (n,m,...) arrays.  In homework/tutorial 1, we created arrays that were 3-d arrays sampling space in a **cube**.  Let's try to cement our intuition about meshgrid by creating a square 2D array and examining its outputs. 


```python
# setting up a 2D, square meshgrid

sample_points_square = np.linspace(0,5,num=6)
print("sample_points_square shape: ", sample_points_square.shape)

square_points_x, square_points_y = np.meshgrid(sample_points_square,
                                              sample_points_square)

print("x coordinates of points in square: \n", square_points_x)
print("y coordinates of points in square: \n", square_points_y)
```

Then, we apply (3,3,3) array that samples a meshgrid and examine its contents. It is hard to represent 3D structures on your 2D montior. One way of visualizing the structure of a 3D array as 2D squre slices of our 'cube' which are represented by 2D arrays. 

Finally, we can use a different 1D array input to make a **rectangular prism** instead of a cube. 

```python
# Setting up points in a cube
sample_points_cube = np.linspace(0,2,num=3)
print("sample_points_cube shape: ", sample_points_cube.shape)
cube_points_x, cube_points_y, cube_points_z = np.meshgrid(sample_points_cube, 
                                                          sample_points_cube, 
                                                          sample_points_cube, 
                                                          indexing='ij')
print("x coordinates of points in cube: \n", cube_points_x)
print("\n shape of x coordinates", cube_points_x.shape)

# Setting up points in a rectangular prism
sample_points_prism_x = np.linspace(0,3,num=4)
print("\n sample_points_prism_x: ", sample_points_prism_x.shape)

#notice we are using a different array for the first input
prism_points_x, prism_points_y, prism_points_z = np.meshgrid(sample_points_prism_x, 
                                                             sample_points_cube, 
                                                             sample_points_cube, 
                                                             indexing='ij')
print("\n x coordinates of points in cube: \n", prism_points_x)
print("\n shape of x coordinates", prism_points_x.shape)
```

**Explain:** 
- How many points sample the cube (how might you find this out)? What would you have to change for there to be 64 points sampling the cube? (1 point)
- We can imagine the output of the cube meshgrid as a list of more intuitive 2D arrays. Based on the output of the cube meshgrid, which coordinate is held constant in a given 2D array? What happens if you change the kwarg `indexing='ij'` to `indexing='xy'`? (1 point)


<!-- grade_id: PF_3D_Array_Questionss, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Array operations (2 points)

Oftentimes, we want to perform an operation on every element of an `ndarray` (e.g. raising them to a power, multiply them, add a constant, etc.). Perhaps the most obvious way to accomplish this is to use a for loop (or many for loops) to iterate over each element of the `ndarray`.

Let's write code to cube every element of an input array.  First, let's define the variable, `base_array` using `np.arange`, starting at 1, stopping at 10.5 in steps of 0.5 (recall, `np.arange` will not include the stopping point).

<!-- grade_id: cell-e87192e148072856, points: 0, grade: false, solution: true -->
```python
# Define base_array below
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: PF_arange, points: 1, grade: true, solution: false -->
```python
"""Execute this to test that you have properly defined base_array"""
assert(base_array[0]==1)
```

Next, let us create a new array that contains each element of `base_array` cubed using a for loop.  There are two ways to do this. Let's start with an implementation using a for loop, which I will do for you. In this code, we:
- Create an empty list, `powered_list_from_loop`.  This is where we will collect the value of each `base_array` element raised to the third power.
- Create a `for` loop to iterate over each `base_array` element, which we will call `base_value`.
- In the `for` loop, set the variable `powered_value` equal to `base_value` to the third power. 
- Still in the `for` loop, append the `powered_value` to the list `powered_list_from_loop`.
- After exiting the `for` loop, create an array, `powered_array_from_loop`,  out of the list `powered_list_from_loop`.

```python
# Create an empty array where we will collect the powered values
powered_list_from_loop = []

# Looping over each value in base_array
for base_value in base_array :
    # set powered_value equal to base_value cubed
    powered_value = base_value**3
    
    # Append powered_value to the list
    powered_list_from_loop.append(powered_value)
    
# Create an array from a list    
powered_array_from_loop = np.array(powered_list_from_loop)
    
print(powered_array_from_loop)
```

One very useful feature of numpy is that both python's basic mathematical functions and numpy's built in functions just *work* on ndarray objects of arbitrary dimension. These array operations can be implemented in a single line of code as if you are just performing the operation on a single value, and for larger datasets can be signficantly faster!

Now, **define a variable**, `powered_array` that is the cube of each value in `base_array`, but do this in one step.  Recall, raising something to a power uses two asterisks, \*\*

<!-- grade_id: cell-50e1afd860f0a3e6, points: 0, grade: false, solution: true -->
```python
# Define powered_array here

# YOUR CODE HERE
raise NotImplementedError()

print(powered_array)
```

<!-- grade_id: PF_cubed_array, points: 1, grade: true, solution: false -->
```python
"""Execute this cell to check that your solution is on the right track"""
assert(powered_array[0] == powered_array_from_loop[0])
```

Here are some other [array operations](https://problemsolvingwithpython.com/05-NumPy-and-Arrays/05.07-Array-Opperations/) you can do.  Output the base 10 log of of `powered_array` in one line. (Note: the last element is 3, because we just did $10^3$, then took the base 10 log of that.)

```python
# Output the base 10 log here

np.log10(powered_array)

```

## Calculating the electric field in a vectorized fashion (3 points)

In the first tutorial/hw, we calculated the electric field at a single point in the field due to a point charge.  To find the electric field at all points we sampled in the field (i.e. the meshgrid), we ended up doing a big loop where we effectively found the euclidean distance between the charge position and each point in the meshgrid.  

```
### THIS IS WHAT WE DID IN TUTORIAL/HW 1

# Create an empty list to collect the electric field at each field position
efield_list = []

charge_position = np.array([0,0,0])

# Loop over all points in the meshgrid
for x, y, z in zip(np.ravel(xarray), np.ravel(yarray), np.ravel(zarray)) :
    field_position = np.array([x,y,z])
    print("Calculating efield at: ", field_position)
    
    # Calculate the electric field at each point
    efield_at_point = calculate_efield_at_point(charge_position, field_position, q=1) 
    print("efield is: ", efield_at_point)
    
    # Append the electric field to a list
    efield_list.append(efield_at_point)
    
efield_vectors = np.array(efield_list)
    
```

Note the parallels between this and the earlier exercise of the preflight.  We will do a similar thing, but will use the `apply_along_axis` function of numpy to do this *all in one line*.   

First, **define your meshgrid points** the same way we did in tutorial/hw 1.

```
num_points_1d = 2
sample_points = np.linspace(-1, 1, num=num_points_1d)
xarray, yarray, zarray = np.meshgrid(sample_points, sample_points, sample_points, indexing='ij')
```

<!-- grade_id: cell-0c0bb8b1995fda5c, points: 0, grade: false, solution: true -->
```python
#  Define your meshgrid points here
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: PF_efield_define_array, points: 1, grade: true, solution: false -->
```python
"""Execute this cell to check that your sample points and meshgrid points are defined correctly"""
assert((sample_points == np.linspace(-1,1,num_points_1d)).all())
assert(xarray[0][0][0]==-1)
assert(yarray[0][0][0]==-1)
assert(zarray[0][0][0]==-1)

```

Let us now create one big array with all points in the meshgrid:
```
points_in_meshgrid = np.array([xarray, yarray, zarray])
```
Print out the shape of this new array that you created.

<!-- grade_id: cell-930ddee173aa1aa3, points: 0, grade: false, solution: true -->
```python
# Define points_in_meshgrid below
# YOUR CODE HERE
raise NotImplementedError()
print(points_in_meshgrid.shape)
```

<!-- grade_id: PF_points_in_meshgrid, points: 1, grade: true, solution: false -->
```python
"""Execute this cell to check that you've properly defined points_in_meshgrid"""
assert(points_in_meshgrid.shape == (3,num_points_1d,num_points_1d,num_points_1d))
```

You'll notice that this is now a 4-d array.  The first dimension corresponds to the number of axes (i.e. x, y, and z means there are three axes). The next 3 dimensions correspond to values of x, y, and z coordinates.  Equivalently, we could have directly defined `points_in_meshgrid`,
```
direct_points_in_meshgrid = np.array(np.meshgrid(sample_points, sample_points, sample_points, indexing='ij'))
```
Define the above in the cell below, and print its shape.

<!-- grade_id: cell-d38dfe3fec03bfe2, points: 0, grade: false, solution: true -->
```python
#  Define direct_points_in_meshgrid here, and print the shape
# YOUR CODE HERE
raise NotImplementedError()
print(direct_points_in_meshgrid.shape)
```

<!-- grade_id: PF_direct_points_in_meshgrid, points: 1, grade: true, solution: false -->
```python
"""Execute this cell to check that you have done the above cell correctly"""
assert((direct_points_in_meshgrid[0] == xarray).all() )
```

Next, we use the same function from the tutorial/hw to calculate the efield at a single field point, but we switch `field_position` with `charge_position` in the order of function arguments.  This is because field position is what gets varied over the entire `points_in_meshgrid`, and we want that to be the first argument of `calculate_efield_at_point`.  Also, let us set `k=1` to simplify visualizations in plots.

```python
def calculate_efield_at_point(field_position, charge_position, q=1) :
    '''Return the electric field due to a point charge.
    
    Inputs:
    field_position (n-darray) : x, y, and z position vector of a field point
    charge_position (n-darray) : x, y, and z position vector of a charge point
    q (float or int) : charge of the point
    Outputs:
    vector_efield (n-darray) : x, y, z components of the e-field at the point field_position
    '''
    
    k = 1 # A choice to simplify plots
    
    r = field_position-charge_position
    r_magnitude = np.linalg.norm(field_position-charge_position)
    r_unit = r/r_magnitude
    
    return k*q / r_magnitude**2 * r_unit
```

With a `charge_position` at the origin, we now calculate `efield_vectors` all in a single line with `apply_along_axis`, see [documentation](https://numpy.org/doc/stable/reference/generated/numpy.apply_along_axis.html),  
```
efield_vectors = np.apply_along_axis(calculate_efield_at_point, 0, points_in_meshgrid, charge_position)
```
the second argument corresponds to the axis along which we apply `calculate_efield_at_point`.  However, our resulting array will have a different shape than when we did it in the loop (n,n,n,3) vs. (3,n,n,n). We will have to take this into account.

<!-- grade_id: cell-084f73b178a18f7b, points: 0, grade: false, solution: true -->
```python
charge_position = np.array([0,0,0])

# Define efield_vectors below
# YOUR CODE HERE
raise NotImplementedError()
print(efield_vectors.shape)
```

<!-- grade_id: PF_apply_along_axis_check, points: 1, grade: true, solution: false -->
```python
"""Execute this to check that you've run the correct code"""
assert(efield_vectors.shape == (3,num_points_1d,num_points_1d,num_points_1d))
efield_vectors[0].ravel().shape
xarray.ravel().shape
```

##  Calculate the efield magnitude (1 point)
Similar to the tutorial/hw 1, we calculate the `efield_magnitude` using `np.linalg.norm`.  We will later use this in a visualization.  First, we have to account for the different shapes.  In the loop, the x, y, z axes varied over the last dimension.  Now, it varies over the first dimension, so the shape is (3,n,n,n).  Note, the axis we take the magnitude over is now the 0th axis (i.e. an $x^2+y^2+z^2$ over each of the points).  So, we will want to calculate the magnitude (norm) of efield_vectors, across the 0th axis,
```
np.linalg.norm(efield_vectors, axis=0)
```

<!-- grade_id: cell-d43e919f5409d14c, points: 0, grade: false, solution: true -->
```python
# Define the efield_magnitude in this cell.
print("Vectors have shape", efield_vectors.shape)

# YOUR CODE HERE
raise NotImplementedError()

print("Magnitude has shape", efield_magnitude.shape)
```

<!-- grade_id: PF_efield_magnitude, points: 1, grade: true, solution: false -->
```python
"""Execute this cell to check that your efield_magnitude has been properly defined"""
print(efield_vectors[:,0,0,0])
assert(efield_magnitude[0][0][0] == np.linalg.norm(efield_vectors[:,0,0,0], axis=0))
```

## Visualize with quiver (1 point)

Similar to tutorial/hw1, we visualize the vector field with quiver.  Let's ravel everything (ravel the xarray, yarray, efield_vectors[0], efield_vectors[1]) so we don't have to worry about aligned shape.  You can do this since `ravel` is a *method* of any numpy array, and all of these are numpy arrays. Or, you can use the function `np.ravel`.  See [documentation](https://numpy.org/doc/stable/reference/generated/numpy.ravel.html). 

Then, go back to where we define `num_points_1d` and change that value to 4.  Run all cells after.

<!-- grade_id: PF_quiver_plot, points: 1, grade: true, solution: true -->
```python
# Use quiver here
#  Plot the vector field using the quiver method
fig, ax1 = plt.subplots(1, figsize=(8,8))

# YOUR CODE HERE
raise NotImplementedError()

ax1.set_xlabel('x (meters)')
ax1.set_ylabel('y (meters)')
ax1.set_title('E Field of Point Charge', fontsize=16)
```