# Phys 260: Python assignment header

### (1) Fill out the cell below.  
The cell below is a **code cell**.  Fill out your University of Michigan uniqname, then your name, and collaborators in the cell below **inside the quotes**.  

**Do not delete the quotes.**  We will use this information to organize your assignments.  To edit and execute cells, double click inside the cell, type, and press \<shift\>+\<enter\> to execute.

```python
UNIQNAME = ""
NAME = ""
COLLABORATORS = ""
```

### (2) Check your python version.  
**Execute the cell below** (double click in the cell and press \<shift\>+\<enter\>, or click in the cell and press the Run button) to check that you are using a version of python that is compatible with the tool we are using to grade your assignments.  If your ```IPython``` version is too old, we will *not* be able to grade your assignments.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

### (3) Do your best to answer all questions in the assignment.  
To answer questions, **replace** anything that says either
- "YOUR ANSWER HERE" 
- 
```
YOUR CODE HERE
raise NotImplementedError
``` 

with your answer/code.  Cells with either of the two bullet points above are cells of the notebook that will be graded.

**To edit markdown** cells (e.g. this one),  *double click in the cell to type*.  Press \<shift\>+\<enter\> to execute the cell.  Try editing the text below to replace the with your information (e.g. Camille Avestruz, cavestru):  

[first name] [last name], uniqname


### (4) Make sure your notebook runs sequentially.
After you complete this assignment, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

```python

```

# Phys 260 Python Lab 3: Numpy tools in review (10 points)

# Introduction

Each Python lab will start with a pre-flight exercise that walks through building some of the set up and tools ($\sim$ 30 min), followed by an in-class tutorial with time for Q+A (50 min) so you can walk through steps that will be necessary for the homework assignment you will submit ($\sim$ 3 hrs).  Each lab will contain starter code, similar to what you see below.  Please fill in the code to complete the pre-flight assignment in preparation for the in-class tutorial.  

Preflight ($\sim$30-60 min, 10 points) **Typically due: Thursdays 3pm ET**

*Preflight typically graded by Wednesday 5p EST -- see your feedback in html (launch a browser)*

In-class tutorial and Q+A ($\sim$ 50 min, 10 points) **Typically occurs: Fridays Noon**

Homework assignment ($\sim$ 3-5 hrs, 30 points) **Typically due: Tuesday Noon**

*Homework typically graded by Thursday 5p -- see your feedback in html (launch a browser)*

When we grade your homework, we will not run your code. Once submitted, your notebook should have the outputs for all of your results.  Please do not include long outputs from debugging, beyond a few print statements and the requested visualizations (i.e. plots).

**Grading:** When we grade your notebook, we will convert the .ipynb file to an HTML file.  We will be using [nbgrader](https://nbgrader.readthedocs.io/en/stable/) to grade your notebooks.  

## Preflight summary
- Brief review of actions along axes
- Useful numpy tools for derivatives
- Effects of discretization: Example with Gauss' Law

```python
# Import relevant modules
import numpy as np
from matplotlib import pyplot as plt
```

## Brief review of actions along axes (5 points)

Below, we review the effects of doing things "along axes"
- With `np.linalg.norm`, 
- with `np.sum`,
- with `np.apply_along_axis`.

First, we use `np.reshape` to create a (4,5) array and examine how this function works.

```python
test_array = np.arange(20)
reshaped_array = np.reshape(test_array, (4,5))
```

**Practice inspection** (1 point):
Inspect both `test_array` and `reshaped_array` (you can print both, and/or print both sizes in the cell above).  What does `np.reshape` do? 

<!-- grade_id: PF_reshape_question, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Norm of different axes** (1 point): 

- Use `np.reshape` with `test_array` to create a `(2,10)` array called `reshaped_2_10`.  Assume these are (x,y) coordinates.  Use `np.linalg.norm` to find the 2-d radius of each of these coordinates, called `radii_2_10`.

- Use Use `np.reshape` with `test_array` to create a `(10,2)` array called `reshaped_10_2`.  Assume these are (x,y) coordinates.  Use `np.linalg.norm` to find the 2-d radius of each of these coordinates, called `radii_10_2`.

<!-- grade_id: cell-5d8ceee85d3fe403, points: 0, grade: false, solution: true -->
```python
# Define reshaped_2_10, radii_2_10, reshaped_10_2, radii_10_2 here
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: PF_reshape_radii, points: 1, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert(radii_10_2.shape == radii_2_10.shape)
assert((reshaped_2_10 == np.reshape(test_array, (2,10))).all())
```

**Question for inspection** (1 point): Are `radii_10_2` and `radii_2_10` the same as one another?  How would you check? (*Hint*: some of the assertion tests show how to do this.)  Why or why are they not the same.

<!-- grade_id: PF_radii_comparison, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Quick exercise** (1 point): Use the appropriately reshaped array from the previous exercise (hint: one of the `reshaped_?_?` arrays) and the function `np.sum` with the appropriate key word argument to return a (2,) array, `sum_even_odd`, where one element corresponds to the sum of all even integers less than 20, and the other element corresponds to all odd integers less than 20.

<!-- grade_id: cell-585b25eff6f56cef, points: 0, grade: false, solution: true -->
```python
# Define sum_even_odd here
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: PF_sum_even_odd, points: 1, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert(sum_even_odd.shape == (2,))
assert(sum_even_odd[0] == np.sum(np.arange(0,20,2)))
assert(sum_even_odd[1] == np.sum(np.arange(1,21,2)))
```

**Quick exercise** (1 point):  Use `np.apply_along_axis` to calculate the same array as `sum_even_odd`, now called `sum_even_odd_from_axis`.  Note, here, you will also need to use `np.sum`, but will not use the key word `axis=`.

<!-- grade_id: cell-65f68f1f9fdff8cf, points: 0, grade: false, solution: true -->
```python
# Define sum_even_odd_from_axis here
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: PF_sum_even_odd_from_axis, points: 1, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert(sum_even_odd_from_axis.shape == (2,))
```

## Useful numpy tools for derivatives

In the next tutorial/hw, you will need to numerically evaluate derivatives.   The derivative is defined as $$\frac {df}{dx} = \lim_{\epsilon \to 0} \frac {f(x+\epsilon) - f(x)}{\epsilon}.$$  To evaulate this numerically we make $\epsilon$ small but finite.  For example if we are calculating a function on a grid of spacing $\delta$  we can approximate the derivative as $$\frac {df}{dx} \approx \frac {f(x+\delta) - f(x)}{\delta}$$  However, we can do a better job by making the derivative symmetric about the point at which we are calculating it:
$$\frac {df}{dx} \approx \frac {f(x+\delta) - f(x - \delta)}{2 \delta}$$ 
To calculate derivatives elements to the right and left of an element.   Array operations give us an easy way to do this.


- `np.roll` applied to taking the numerical derivative (see how it works!)
- `np.gradient` pretty much what you were doing before, but with a built-in function

Let's say our function is $f=x^2$, and x is our `test_array` so each element corresponds to our spacign where $\delta=1$.  First, let's find the $f(x+\delta)$, which we will call `f_right`:

```python
f = test_array**2
f_right = np.roll(f,-1)
```

**Quick exercise** (1 point): Use `np.roll` (examine what happened with `f_right` above, and see the [documentation](https://numpy.org/doc/stable/reference/generated/numpy.roll.html)) to find `f_left` ($f(x-\delta)$), and calculate `df_dx`.

<!-- grade_id: cell-1f7dbfc9406d7de9, points: 0, grade: false, solution: true -->
```python
# Calculate f_left and df_dx here
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: PF_df_dx, points: 1, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert((f_left == np.roll(f,1)).all())
```

**Plot to check** (1 point) You'll likely notice that the boundaries (i.e. first and last elements) look funky.  With general simulations of physical systems, one often creates boundary conditions (e.g. periodic boundary conditions) or ignores the results in the boundary.  Plot `df_dx` vs. `test_array`, ignoring the first and last elements.  You'll need to recall array slicing to do that (*Hint*: the first element is indexed by 0, so you'll want to start at index 1, and the last element is indexed by -1).  Does this look like what you expect given the derivative of $x^2$?  Full credit with labeled axes and legible axes (I like to recomment 'xx-large' fontsize for axes labels).

<!-- grade_id: PF_plot_derivative, points: 1, grade: true, solution: true -->
```python
# Plot df_dx here
# YOUR CODE HERE
raise NotImplementedError()
```

**Quick exercise** (1 point):  Now, find `df_dx_gradient` using `np.gradient` (see [documentation](https://numpy.org/doc/stable/reference/generated/numpy.gradient.html)).

<!-- grade_id: cell-a8bb560e32057ea2, points: 0, grade: false, solution: true -->
```python
# Find df_dx_gradient here
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: PF_np_gradient, points: 1, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
assert(df_dx_gradient.shape == f.shape)
```

**Quick comparison** (1 point):  You'll notice a key difference between doing the `np.roll` then `np.difference` with doing the `np.gradient`.  Describe the difference in output.  (*Hint:* The difference is because there's an under-the-hood treatment for edges with `np.gradient`.)

<!-- grade_id: PF_np_gradient_comparison, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Overplot** (1 point):  Overplot the gradient of f using the `kwargs` `edge_order=1` and `edge_order=2` in `np.gradient`.  Label these with "edge order 1" and "edge order 2", create a legend.  Observe the difference between first order accurate difference and second order accurate difference (you can think of this as similar to taking the Taylor expansion to first or second order... the further out you go, the more accurate the approximation).  First order essentially linearly extrapolates to the left and right of the edges of f to calculate the derivative at the edges, and second order essentially fits a second order (parabola) through three points on the left and on the right to extrapolate and calculate the derivative.

<!-- grade_id: PF_plot_Gradient_edge_order, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```