# Phys 260: Python assignment header

### (1) Fill out the cell below.  
The cell below is a **code cell**.  Fill out your University of Michigan uniqname, then your name, and collaborators in the cell below **inside the quotes**.  

**Do not delete the quotes.**  We will use this information to organize your assignments.  To edit and execute cells, double click inside the cell, type, and press \<shift\>+\<enter\> to execute.

```python
UNIQNAME = ""
NAME = ""
COLLABORATORS = ""
```

### (2) Check your python version.  
**Execute the cell below** (double click in the cell and press \<shift\>+\<enter\>, or click in the cell and press the Run button) to check that you are using a version of python that is compatible with the tool we are using to grade your assignments.  If your ```IPython``` version is too old, we will *not* be able to grade your assignments.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

### (3) Do your best to answer all questions in the assignment.  
To answer questions, **replace** anything that says either
- "YOUR ANSWER HERE" 
- 
```
YOUR CODE HERE
raise NotImplementedError
``` 

with your answer/code.  Cells with either of the two bullet points above are cells of the notebook that will be graded.

**To edit markdown** cells (e.g. this one),  *double click in the cell to type*.  Press \<shift\>+\<enter\> to execute the cell.  Try editing the text below to replace the with your information (e.g. Camille Avestruz, cavestru):  

[first name] [last name], uniqname


### (4) Make sure your notebook runs sequentially.
After you complete this assignment, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

```python

```

# Introduction -- Reminder

Each Python lab will start with a pre-flight exercise that walks through building some of the set up and tools ($\sim$ 30 min), followed by an in-class tutorial with time for Q+A (50 min) so you can walk through steps that will be necessary for the homework assignment you will submit ($\sim$ 3 hrs).  Each lab will contain starter code, similar to what you see below.  Please fill in the code to complete the pre-flight assignment in preparation for the in-class tutorial.  

Preflight ($\sim$30-60 min, 10 points) **Typically due: Wednesdays 3pm EST**

*Preflight typically graded by Wednesday 5p EST*

In-class tutorial and Q+A ($\sim$ 50 min, 10 points) **Typically occurs: Thursdays 9am EST**

Homework assignment ($\sim$ 3-5 hrs, 30 points) **Typically due: Mondays 9am EST** 

*Homework typically graded by Thursday 5p*

When we grade your homework, we will not run your code. Once submitted, your notebook should have the outputs for all of your results.  Please do not include long outputs from debugging, beyond a few print statements and the requested visualizations (i.e. plots).

**Grading:** When we grade your notebook, we will convert the .ipynb file to an HTML file.  We will be using [nbgrader](https://nbgrader.readthedocs.io/en/stable/) to grade your notebooks.  **Note:** If your ```IPython``` version is too old, we will *not* be able to grade your assignments.


# Phys 260 Python Preflight 9: Simple circuits with inductors  (10 points total)

## Preflight Summary
- Model of a circuit with R, L and C components

In this assignment we will examine the effect of inductors in circuits.  For the tutorial/hw, we will introduce the concept of alternating-current circuits. 

The canonical example of an inductor is a solenoid. When you run a current through a solenoid, it develops a magnetic field. If this current is changing, the magnetic field in the solenoid is changing too. And if that happens, the solenoid experiences a changing magnetic flux through itself. Faraday's law tells us that this will result in an emf in the solenoid that's in a direction to oppose the change in flux. 

The induced emf in the solenoid is proportional to the rate of change of the current, so we can write it as:
$$
    V_L = -L\frac{dI}{dt}.
$$
$L$ is a proportionality constant called the <b>inductance</b>. It's measured in Henries. 

Like capacitance, inductance is a geometrical quantity. We'll calculate it explicitly for a solenoid in class. Physically, inductors prevent currents from changing too quickly. If $dI/dt$ is large, the inductor will generate a large opposing emf that limits the current. Surge protectors are basically big inductors.

Inductors contribute to Kirchhoff's loop rule, just like resistors, capacitors, and batteries do. So we can solve circuits with inductors using techniques similar to what you did in the previous preflight.

```python
import numpy as np
from matplotlib import pyplot as plt
```

## Part 1 (10 points): A simple RL circuit

Consider this circuit:

<img src="http://www-personal.umich.edu/~gerdes/img/LR_DC.jpg">


### LaTeX the solution (4 points)
At $t=0$ the switch is closed.  Use LaTeX in the cell below to: 
- Write the loop rule (sum of voltage drops) for this circuit,
- Write the solution for the differential equation, $I(t)$, in terms of $V_{\rm batt}$, R, and L,
- Write the expression for the time constant, $\tau$,
- Write the voltage drop across the inductor as a function of time, $V_{L}$, in terms of $V_{\rm batt}$ and $\tau$.  Note, $V_{L}=-L\frac{dI}{dt}$.  This will be useful in modeling the current evolution.

<!-- grade_id: cell-01406401414c8319, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

### Create a class to model the RL circuit (2 points)

As a check, your LaTeX equations should show that the current starts at zero and exponentially approaches its asymptotic value of $V/R$ (as the inductor approaches the behavior of a wire).  The inductor voltage starts out at $V$ and exponentially decays to zero as the current becomes steady (no change in current means no voltage drop across the inductor).

In the cell below, create a class to model the RL circuit.  You may want to refer to the structure from the previous preflight.  Some differences include the fact that there is only one current, the initial current is zero (e.g. the `self.i_now`), and the characteristic timescale `self.tau` is different.

One recommendation in updating the now value step is to calculate the voltage drop across the inductor, given the current, then update the current 

<!-- grade_id: cell-c09b769a8d55e1d4, points: 0, grade: false, solution: true -->
```python
class RLCircuit :
    def __init__(self,voltage_battery, resistance, inductance) :
        """
        Parameters
        ----------
        voltage_battery : float
            voltage across battery
        resistance : float
            value of resistance in ohms
        inductance : float 
            value of inductance in Henries
        """
        
        self.voltage_battery = voltage_battery
        self.resistance = resistance
        self.inductance = inductance
        
        # Initialize other necessary attributes 
        # YOUR CODE HERE
        raise NotImplementedError()

        
    def _set_now_vals(self) :
        """Calculate "now" current in terms of other attributes, setting attributes"""
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def _set_timesteps(self, num_timescales, num_timesteps_per_timescale) :        
        """Internal method. Sets the attributes num_timesteps and dt.
        Parameters
        ----------
        num_timescales : floattimestep 
            number of timescales to evolve over
        num_timesteps_per_timescale : int
            number of timesteps per characteristic timescale, defines the time resolution 
        """
        
        self.dt = self.tau / num_timesteps_per_timescale
        print(self.dt)

        next_timesteps = np.arange(0, self.tau * num_timescales, self.dt)
        print('next_timesteps shape: ',next_timesteps.shape)
        try :
            last_timestep = self.timesteps[-1]
            next_timesteps += last_timestep
            self.timesteps = np.concatenate([self.timesteps, next_timesteps])
            
        except IndexError : 
            self.timesteps = next_timesteps        

    def evolve_circuit(self, num_timescales, num_timesteps_per_timescale) :
        """Evolve the spring, populate the acceleration, velocity and position (a_vals, v_vals, x_vals)
        Parameters
        ----------
        num_periods : float
            number of periods to evolve over
        num_timesteps_per_period : int
            number of timesteps per period, defines the time resolution         
        """
        
        self._set_timesteps(num_timescales, num_timesteps_per_timescale)

        # Time evolve the circuit using self._set_now_vals() and store values at each timestep
        # YOUR CODE HERE
        raise NotImplementedError()
```

<!-- grade_id: cell-a41d7f626ccbd6c1, points: 2, grade: true, solution: false -->
```python
"""Execute to check you're on the right track"""
test_rl = RLCircuit(10,1,1)
test_rl.evolve_circuit(5,100)
assert(test_rl.i_vals[0] < test_rl.i_vals[-1])
```

### Create an instance and plot (2 points)
Create an instance of `RLCircuit` with the following components:
* V = 10 Volts
* R = 1 $\Omega$
* L = 1 H

Time evolve and plot:
* On the same axis: Voltage across the resistor vs. time (in units of $\tau$) and Voltage across the inductor vs. time.  Note, you may find it useful to store the voltage across the inductor at each timestep, but there are multiple ways to do this.
* On the same axis: Power delivered to the resistor vs. time and power delivered to the inductor vs. time (in units of $\tau$). Recall that $P = IV$. 

<!-- grade_id: cell-8b2ef77437a9350f, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Open questions (2 points) :

<b>Answer</b> the following questions:
* At what time (in units of $\tau$) is the maximum power being delivered to the inductor? 
* What is the total energy $U$ delivered to the inductor? Recall, when discretizing, power at a given timestep is energy per unit time.  Multiplying by the timestep size gives the energy over a single timestep.  
* If you change $L$ while holding $R$ constant, how does $U$ change? 
* If you change $R$ while holding $L$ constant, how does $U$ change?
* If you simultaneously double $R$ and $L$, you get a circuit with the same time constant as before. Is this new circuit equivalent to your original one?

<!-- grade_id: cell-b55e989d979c7901, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE